## MISA (2024-2025)
- Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

- Izay misy hoe `YOUR CODE HERE` na `YOUR ANSWER HERE` ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
Eto ilay references rehetra no apetraka

---

```python
from sklearn.datasets import load_diabetes
import numpy as np
from sklearn.kernel_ridge import KernelRidge

def rel_error(x, y):
  """ returns relative error """
  return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))

data = load_diabetes()
X_train, y_train = data.data, data.target
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 1
----> 1 from sklearn.datasets import load_diabetes
      2 import numpy as np
      3 from sklearn.kernel_ridge import KernelRidge

ModuleNotFoundError: No module named 'sklearn'
```

<!-- grade_id: cell-d36cbb79389907e8, points: 0, grade: false, solution: true -->
```python
def gaussian_kernel(x1, x2, sigma=1.0):
    x1 = np.asarray(x1)
    x2 = np.asarray(x2)
    x1_2d = np.atleast_2d(x1)
    x2_2d = np.atleast_2d(x2)
    diff = x1_2d[:, None, :] - x2_2d[None, :, :]
    sq_dist = np.sum(diff ** 2, axis=2)
    k = np.exp(-sq_dist / (2.0 * sigma ** 2))
    if x1.ndim == 1 and x2.ndim == 1:
        return float(k[0, 0])
    return k

class KernelRidgeRegression():
    def __init__(self, kernel, lambd=1.0):
        self.lambd = lambd # regularization coefficient
        self.alpha = None # dual variable
        self.kernel = kernel # kernel function
        self.X_train = None

    def fit(self, X, y):
        self.X_train = np.asarray(X)
        y = np.asarray(y)
        K = self.kernel(self.X_train, self.X_train)
        n_samples = K.shape[0]
        K_reg = K + self.lambd * np.eye(n_samples)
        self.alpha = np.linalg.solve(K_reg, y)

    def predict(self, X):
        X = np.asarray(X)
        K = self.kernel(self.X_train, X)
        return K.T @ self.alpha
```

<!-- grade_id: cell-ec40f17b949370cd, points: 5, grade: true, solution: false -->
```python
kr = KernelRidgeRegression(kernel=gaussian_kernel)
kr.fit(X_train,y_train)
y_pred = kr.predict(X_train)

sklearn_kr = KernelRidge(kernel=gaussian_kernel)
sklearn_kr.fit(X_train,y_train)
sklearn_y_pred = kr.predict(X_train)

error = rel_error(kr.alpha, sklearn_kr.dual_coef_)
print("error: ", error)
assert error <= 1e-11

error_pred = rel_error(y_pred, sklearn_y_pred)
print("prediction error: ", error_pred)
assert error_pred <= 1e-11
```